# Working with a database

This week's assignment has a few basic steps.  First, we're going to pull some data down off the internet and store it into our MySQL database.  Make sure that you use your username as part of the table name as show in the examples so that you don't create a problem for other students.

Then, we'll merge that with some data already in the database and calculate a few results.  When it comes to calculating the results, you can do so either with SQL or with Pandas operations.


## PART 1: Setup your database connection and table name

In the code below, change the value of the variable `MYTABLE` to use your own username rather that `'pboal'`

You can then use `MYTABLE` in the rest of your code to reference that table name.

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!
# First, put in your user name below
#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!
import os
import pymysql
from sqlalchemy import create_engine
import pandas as pd
host = 'slucor2020.cgdcoitnku0k.us-east-1.rds.amazonaws.com'
port = '3306'
user = 'slucor2020'
password = 'SLUcor2020'
database = 'hds5210'

MYTABLE='kenny0415' + '_data'


# Then, create a connection to the same database, using the 
# same user name and password in the `mysql-examples.ipynb` 
# file.  Use conn as the database connection variable.

conn =  create_engine('mysql+pymysql://' + 
                     user + ':' + 
                     password + '@' + 
                     host + '/' + 
                     database, echo=False)

```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
assert(MYTABLE != 'paulboal_data')
assert(conn.connect())
```

## PART 2: Bring in outside data

Grab the data from this URL and put it into a database table named with your `username_data`.

http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv

By the end of your cell, the table should be created.  The tests are going to verify that the table exists and looks right.

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
url = 'http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv'
df=pd.read_csv(url)
df.to_sql(MYTABLE, conn, index=False)

### SOLUTION GOES HERE

```

```python
dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
dxyz.shape
```

Output:
```
(328, 19)
```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python
dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
assert(dxyz.shape == (328,19))
assert(list(dxyz.columns) == ['Provider_Name', 'NPI', 'CCN',
       'Medicaid_EP_Hospital_Type', 'Street_Address', 'City', 'County',
       'State', 'Zip_Code', 'Payment_Year_Number', 'Program_Type',
       'Total_payments', 'Last_Program_Year', 'Last_Payment_Year',
       'Last_Payment_Criteria', 'Most_Recent_Disbursement_Amount', 'LONGITUDE',
       'LATITUDE', 'FID'])
```

```python
#with conn.connect() as c:
     #c.execute('DROP TABLE IF EXISTS ' +MYTABLE)
```

## PART 3: Combine with other data in the database

In the database, there is an existing table called `population`.  We want to merge the DHCS datafile loaded above with the population data available in this other database table  The tables can be merged on `MYTABLE`'s `Zip_Code` field and `population`'s `zipcode` field.

Note that not all `Zip_Codes` from your downloaded file have to be in the `population` table.  If they aren't, then I want you to eliminate the non-matching records.  That is, only keep the records that have a matching ZIP code in both sets of data.

Answer the question:
Which providers are located in the zipcode with the largest population?

Put your answer in the form `answer = ['a', 'list', 'of', 'NPI', 'like', '1593042103]`

```python
pd.read_sql_query('SHOW TABLES', conn)
```

Output:
```
Tables_in_hds5210
0                      NPI
1            Provider_Name
2                     Test
3              aavila_data
4          aavila_data_NPI
5          aavila_data_new
6         aavila_data_year
7             aavila_datax
8                 abbieslu
9             abbieslu_NPI
10            abbieslu_ipv
11            abbieslu_npi
12           abbieslu_year
13               achenliam
14          achenliam_data
15           achenliam_ipv
16           apuryear_data
17  arefehesmaeilpour_data
18             ashley_data
19            btabman_data
20           corona_counts
21       county_population
22       dhumphrey420_data
23               ed_visits
24       hospital_services
25               hospitals
26          jdgendron_data
27           jmccar23_data
28       kaylaschmidt_data
29          kenny0415_data
30            khalafh_data
31          lchen_data_ipv
32         liannecota_data
33           paulboal_data
34              population
35                ppp_data
36        
```

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
### SOLUTION GOES HERE
answer=pd.read_sql_query("""
  SELECT k.NPI
  FROM 
    population p INNER JOIN
    kenny0415_data k ON k.Zip_Code = p.zipcode
    ORDER BY 
    p.population desc limit 1;
""", conn)
#answer=answer.astype(str).values.tolist()[0]
list(answer['NPI'].astype(str))
```

Output:
```
['1194016923']
```

```python
pd.read_sql_query("SELECT * FROM population", conn)
```

Output:
```
zipcode  population  median_age  total_males  total_females  \
0     91371           1        73.5            0              1   
1     90001       57110        26.6        28468          28642   
2     90002       51223        25.5        24876          26347   
3     90003       66266        26.3        32631          33635   
4     90004       62180        34.8        31302          30878   
..      ...         ...         ...          ...            ...   
314   93552       38158        28.4        18711          19447   
315   93553        2138        43.3         1121           1017   
316   93560       18910        32.4         9491           9419   
317   93563         388        44.5          263            125   
318   93591        7285        30.9         3653           3632   

     total_households  avg_household  
0                   1           1.00  
1               12971           4.40  
2               11731           4.36  
3               15642           4.22  
4   
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == list)
assert(answer == ['1194016923'])
```

## PART 4: Total by hospital type

This is a multistep process:
* From our downloaded data file, compute the `Total payments` per ZIP code and Medicaid EP Hospital Type.
* Then merge that with the `population` data to compute a `Total payments` per person.
* Then average that across all of the `Medicaid EP Hospital Types` to get an average per persona payment for these type of hospital.

Your answer should be in structure of a data frame with at least two columns:
* Medicaid_EP_Hospital_Type
* Avg_Pay_per_Capita

```python
tables = {}
for n,t in pd.read_sql_query("SHOW TABLES", conn).iterrows():
    name = str(t['Tables_in_hds5210'])
    print(name)
    tables[name] = pd.read_sql("SELECT * FROM "+name,conn)
```

Output:
```
NPI
Provider_Name
Test
aavila_data
aavila_data_NPI
aavila_data_new
aavila_data_year
aavila_datax
abbieslu
abbieslu_NPI
abbieslu_ipv
abbieslu_npi
abbieslu_year
apuryear_data
arefehesmaeilpour_data
ashley_data
btabman_data
corona_counts
county_population
ed_visits
hospital_services
hospitals
jdgendron_data
jmccar23_data
kaylaschmidt_data
kenny0415_data
khalafh_data
liannecota_data
paulboal_data
population
sam
sampele
system_affiliations
x
```

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
### SOLUTION GOES HERE
tables['kenny0415_data']['Zip_Code']=tables['kenny0415_data']['Zip_Code'].astype(str)
kenny0415_data_groupby=tables['kenny0415_data'].groupby(["Zip_Code","Medicaid_EP_Hospital_Type"]).sum().reset_index()
tables['population']['zipcode']=tables['population']['zipcode'].astype(str)
merge=kenny0415_data_groupby.merge(tables['population'], how='inner', left_on='Zip_Code', right_on='zipcode')
merge['Total_payments_per_person']=merge['Total_payments'] / merge['population']
merge=merge.groupby('Medicaid_EP_Hospital_Type').mean()
merge['Avg_Pay_per_Capita']=merge['Total_payments_per_person']
answer = merge
answer
```

Output:
```
NPI           CCN  Payment_Year_Number  \
Medicaid_EP_Hospital_Type                                                    
Acute Care Hospitals       1.693329e+09  56481.066667             3.186667   
Children's Hospitals       1.370193e+09  53307.000000             2.666667   

                           Total_payments  Last_Program_Year  \
Medicaid_EP_Hospital_Type                                      
Acute Care Hospitals         3.245626e+06        2256.213333   
Children's Hospitals         4.847082e+06        2014.666667   

                           Last_Payment_Year  Most_Recent_Disbursement_Amount  \
Medicaid_EP_Hospital_Type                                                       
Acute Care Hospitals             2257.533333                    760936.680000   
Children's Hospitals             2015.666667                    592707.333333   

                           LONGITUDE   LATITUDE    FID    population  \
Medicaid_EP_Hospital_Type                                            
```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == pd.core.frame.DataFrame)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 17974.363)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 111.559)

```

```python

```

```python

```